Before you turn this notebook in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below:

```python
NAME = ""
```

_Understanding Deep Learning_

---

<a href="https://colab.research.google.com/github/DL4DS/sp2024_notebooks/blob/main/release/2_1_Supervised_Learning.ipynb" target="_blank"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Notebook 2.1 Supervised Learning

The purpose of this notebook is to explore the linear regression model discussed in Chapter 2 of the book.

Work through the cells below, running each cell in turn. In various places you will see the words "YOUR CODE HERE".
Follow the instructions at these places and write code to complete the functions. There are also questions interspersed
in the text. Put your answer where it says "YOUR ANSWER HERE".

```python
# Math library
import numpy as np
import math
# Plotting library
import matplotlib.pyplot as plt
```

```python
# Create some input / output data
x = np.array([0.03, 0.19, 0.34, 0.46, 0.78, 0.81, 1.08, 1.18, 1.39, 1.60, 1.65, 1.90])
y = np.array([0.67, 0.85, 1.05, 1.0, 1.40, 1.5, 1.3, 1.54, 1.55, 1.68, 1.73, 1.6 ])

print(x)
print(y)
```

<!-- grade_id: cell-aeb464eca8ca783a, points: 0, grade: false, solution: true -->
```python
# Define 1D linear regression model
def f(x, phi0, phi1):
  # Implement the linear regression model (eq 2.4)

  # YOUR CODE HERE
  raise NotImplementedError()

  return y
```

<!-- grade_id: cell-39048e1a2417bc33, points: 2, grade: true, solution: false -->
```python
f(0.3, 0.5, 0.5) == 0.65
```

```python
# Function to help plot the data
def plot(x, y, phi0, phi1):
    fig,ax = plt.subplots()
    ax.scatter(x,y)
    plt.xlim([0,2.0])
    plt.ylim([0,2.0])
    ax.set_xlabel('Input, $x$')
    ax.set_ylabel('Output, $y$')
    # Draw line
    x_line = np.arange(0,2,0.01)
    y_line = f(x_line, phi0, phi1)
    plt.plot(x_line, y_line,'b-',lw=2)

    plt.show()
```

```python
# Set the intercept and slope as in figure 2.2b
phi0 = 0.4 ; phi1 = 0.2
# Plot the data and the model
plot(x,y,phi0,phi1)
```

<!-- grade_id: cell-f0a20716861fb74f, points: 0, grade: false, solution: true -->
```python
# Function to calculate the loss
def compute_loss(x,y,phi0,phi1):

  # Implement the loss calculation (equation 2.5)

  # YOUR CODE HERE
  raise NotImplementedError()

  return loss
```

```python
# Compute the loss for our current model
loss = compute_loss(x,y,phi0,phi1)
print(f'Your Loss = {loss:3.2f}, Ground truth =7.07')
```

<!-- grade_id: cell-8fe40058419d1f31, points: 2, grade: true, solution: false -->
```python
assert math.isclose(loss, 7.07, rel_tol=1e-2), "Your loss function is not correct"
```

```python
# Let's set the intercept and slope as in figure 2.2c
phi0 = 1.60 ; phi1 =-0.8
# Plot the data and the model
plot(x,y,phi0,phi1)
loss = compute_loss(x,y,phi0,phi1)
print(f'Your Loss = {loss:3.2f}, Ground truth =10.28')
```

### Problem

Change the parameters below manually to fit the model.

1. First fix phi1 and try changing phi0 until you can't make the loss go down any more
2. Then fix phi0 and try changing phi1 until you can't make the loss go down any more
3. Repeat this process until you find a set of parameters that fit the model as in figure 2.2d

Leave your final values in the cell below.

```python
phi0 = 1.60 ; phi1 =-0.8    # initial values. leave as is.
```

<!-- grade_id: cell-08d220f4732ea845, points: 3, grade: true, solution: true -->
```python
# assign phi0 and phi1 below

# YOUR CODE HERE
raise NotImplementedError()

```

```python
plot(x,y,phi0,phi1)
print(f'Your Loss = {compute_loss(x,y,phi0,phi1):3.2f}')
```

# Visualizing the loss function

The above process is equivalent to to descending coordinate wise on the loss function<br>

Now let's plot that function.

```python
# Make a 2D grid of possible phi0 and phi1 values
phi0_mesh, phi1_mesh = np.meshgrid(np.arange(0.0,2.0,0.02), np.arange(-1.0,1.0,0.02))

# Make a 2D array for the losses
all_losses = np.zeros_like(phi1_mesh)
# Run through each 2D combination of phi0, phi1 and compute loss
for indices,temp in np.ndenumerate(phi1_mesh):
    all_losses[indices] = compute_loss(x,y, phi0_mesh[indices], phi1_mesh[indices])

```

```python
# Plot the loss function as a heatmap
fig = plt.figure()
ax = plt.axes()
fig.set_size_inches(7,7)
levels = 256
ax.contourf(phi0_mesh, phi1_mesh, all_losses ,levels)
levels = 40
ax.contour(phi0_mesh, phi1_mesh, all_losses ,levels, colors=['#80808080'])
ax.set_ylim([1,-1])
ax.set_xlabel('Intercept, $\phi_0$')
ax.set_ylabel('Slope, $\phi_1$')

# Plot the position of your best fitting line on the loss function
# It should be close to the minimum
ax.plot(phi0,phi1,'ro')
plt.show()
```