While editing this notebook, don't change cell types as that confuses the autograder.

Before you turn this notebook in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below:

```python
NAME = ""
```

_Understanding Deep Learning_

---

<a href="https://colab.research.google.com/github/DL4DS/sp2024_notebooks/blob/main/release/nbs10/10_1_1D_Convolution.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Notebook 10.1: 1D Convolution

This notebook investigates 1D convolutional layers.

Adapted from notebooks at https://github.com/udlbook/udlbook.

> Note: A convolutional filter with no spaces between the elements (i.e. a normal filter without dilation) as denoted as having dilation of one.

```python
import numpy as np
import matplotlib.pyplot as plt
```

## Define a Signal

```python
# Define a signal that we can apply convolution to
x = [5.2, 5.3, 5.4, 5.1, 10.1, 10.3, 9.9, 10.3, 3.2, 3.4, 3.3, 3.1]
```

```python
# Draw the signal
fig,ax = plt.subplots()
ax.plot(x, 'k-')
ax.set_xlim(0,11)
ax.set_ylim(0, 12)
plt.show()
```

## A 3/1/1 Convolution

- kernel size: 3
- stride: 1
- dilation: 1

<!-- grade_id: cell-79afd2b9ee91c840, points: 0, grade: false, solution: true -->
```python
# Now let's define a zero-padded convolution operation
# with a convolution kernel size of 3, a stride of 1, and a dilation of 1
# as in figure 10.2a-c.  Write it yourself, don't call a library routine!
# Don't forget that Python arrays are indexed from zero, not from 1 as in the book figures
def conv_3_1_1_zp(x_in, omega):
    x_out = np.zeros_like(x_in)

    # YOUR CODE HERE
    raise NotImplementedError()

    return x_out
```

Now let's see what kind of things convolution can do
First, it can average nearby values, smoothing the function:

```python

omega = [0.33,0.33,0.33]
h = conv_3_1_1_zp(x, omega)

# Check that you have computed this correctly
print(f"Sum of output is {np.sum(h):3.3}, should be 71.1")

# Draw the signal
fig,ax = plt.subplots()
ax.plot(x, 'k-',label='before')
ax.plot(h, 'r-',label='after')
ax.set_xlim(0,11)
ax.set_ylim(0, 12)
ax.legend()
plt.show()
```

<!-- grade_id: cell-d7c4326df6133552, points: 1, grade: true, solution: false -->
```python
assert np.isclose(np.sum(h), 71.1, atol=0.1), "Sum of output is not correct"
```

Notice how the red function is a smoothed version of the black one as it has averaged adjacent values.  The first and last outputs are considerably lower than the original curve though.  Make sure that you understand why!<br><br>

With different weights, the convolution can be used to find sharp changes in the function:

```python
omega = [-0.5,0,0.5]
h2 = conv_3_1_1_zp(x, omega)

# Draw the signal
fig,ax = plt.subplots()
ax.plot(x, 'k-',label='before')
ax.plot(h2, 'r-',label='after')
ax.set_xlim(0,11)
# ax.set_ylim(0, 12)
ax.legend()
plt.show()
```

<!-- grade_id: cell-94d64e00b85a6414, points: 1, grade: true, solution: false -->
```python
assert np.isclose(np.sum(h2), -1.05, atol=0.1), "Sum of output is not correct"
```

Notice that the convolution has a peak where the original function went up and trough where it went down.  It is roughly zero where the function is locally flat.  This convolution approximates a derivative.


## A 3/2/1 Convolution

Now let's define the convolutions from figure 10.3.  

* kernel size: 3
* stride: 2
* dilation: 1

<!-- grade_id: cell-bd8c2f8278f0c12c, points: 0, grade: false, solution: true -->
```python
# Now let's define a zero-padded convolution operation
# with a convolution kernel size of 3, a stride of 2, and a dilation of 1
# as in figure 10.3a-b.  Write it yourself, don't call a library routine!
def conv_3_2_1_zp(x_in, omega):
    x_out = np.zeros(int(np.ceil(len(x_in)/2)))
    
    # YOUR CODE HERE
    raise NotImplementedError()

    return x_out
```

```python
omega = [0.33,0.33,0.33]
h3 = conv_3_2_1_zp(x, omega)

# If you have done this right, the output length should be six and it should
# contain every other value from the original convolution with stride 1
print(h)
print(h3)
```

<!-- grade_id: cell-cffeb297166c2fb2, points: 1, grade: true, solution: false -->
```python
assert np.allclose(h[::2], h3), "Output is not correct"
```

## A 5/1/1 Convolution

* kernel size: 5
* stride: 1
* dilation: 1

<!-- grade_id: cell-43afee3b0c49402f, points: 0, grade: false, solution: true -->
```python
# Now let's define a zero-padded convolution operation
# with a convolution kernel size of 5, a stride of 1, and a dilation of 1
# as in figure 10.3c.  Write it yourself, don't call a library routine!
def conv_5_1_1_zp(x_in, omega):
    x_out = np.zeros_like(x_in)

    # YOUR CODE HERE
    raise NotImplementedError()


    return x_out
```

```python

omega2 = [0.2, 0.2, 0.2, 0.2, 0.2]
h4 = conv_5_1_1_zp(x, omega2)

# Check that you have computed this correctly
print(f"Sum of output is {np.sum(h4):3.3}, should be 69.6")

# Draw the signal
fig,ax = plt.subplots()
ax.plot(x, 'k-',label='before')
ax.plot(h4, 'r-',label='after')
ax.set_xlim(0,11)
ax.set_ylim(0, 12)
ax.legend()
plt.show()
```

<!-- grade_id: cell-a4913e7b1e2cfd42, points: 1, grade: true, solution: false -->
```python
assert np.isclose(np.sum(h4), 69.6, atol=0.1), "Sum of output is not correct"
```

## A 3/1/2 Convolution

* kernel size: 3
* stride: 1
* dilation: 2

<!-- grade_id: cell-62ab3aeff0c088fe, points: 0, grade: false, solution: true -->
```python
# Finally let's define a zero-padded convolution operation
# with a convolution kernel size of 3, a stride of 1, and a dilation of 2
# as in figure 10.3d.  Write it yourself, don't call a library routine!
# Don't forget that Python arrays are indexed from zero, not from 1 as in the book figures
def conv_3_1_2_zp(x_in, omega):
    x_out = np.zeros_like(x_in)

    # YOUR CODE HERE
    raise NotImplementedError()

    return x_out
```

```python
omega = [0.33,0.33,0.33]
h5 = conv_3_1_2_zp(x, omega)

# Check that you have computed this correctly
print(f"Sum of output is {np.sum(h5):3.3}, should be 68.3")

# Draw the signal
fig,ax = plt.subplots()
ax.plot(x, 'k-',label='before')
ax.plot(h5, 'r-',label='after')
ax.set_xlim(0,11)
ax.set_ylim(0, 12)
ax.legend()
plt.show()
```

<!-- grade_id: cell-f6d3188b9ef0bca7, points: 1, grade: true, solution: false -->
```python
assert np.isclose(np.sum(h5), 68.3, atol=0.1), "Sum of output is not correct"
```

## Convolutions as Matrices

Finally, let's investigate representing convolutions as full matrices, and show we get the same answer.

<!-- grade_id: cell-2030de30aef582ff, points: 0, grade: false, solution: true -->
```python
# Compute matrix in figure 10.4 d
def get_conv_mat_3_1_1_zp(n_out, omega):
  omega_mat = np.zeros((n_out,n_out))
  # TODO Fill in this omega_mat with the correct values

  # YOUR CODE HERE
  raise NotImplementedError()

  return omega_mat
```

```python
# Run original convolution
omega = np.array([-1.0,0.5,-0.2])
h6 = conv_3_1_1_zp(x, omega)
print(h6)

# If you have done this right, you should get the same answer
omega_mat = get_conv_mat_3_1_1_zp(len(x), omega)
h7 = np.matmul(omega_mat, x)
print(h7)

```

<!-- grade_id: cell-e0f8d5e81bff598e, points: 1, grade: true, solution: false -->
```python
assert np.allclose(h6, h7), "Output is not correct"
```

```python
print(omega_mat)
```

TODO:  What do you expect to happen if we apply the last convolution twice?  Can this be represented as a single convolution?  If so, then what is it?

<!-- grade_id: cell-7b257dcab2cab9c7, points: 0, grade: false, solution: true -->
```python
# Apply the matrix to x twice
h7 = np.matmul(omega_mat, x)   # recalculate h7
h8 = np.matmul(omega_mat, h7)
print(h8)

# Create an omega_mat2 that when applied to x gives the same results as convolving
# omega_mat with x twice.
omega_mat2 = None
# YOUR CODE HERE
raise NotImplementedError()

h9 = np.matmul(omega_mat2, x)
print(h9)

h6 = conv_3_1_1_zp(x, omega)  # recalculate h6
h10 = conv_3_1_1_zp(h6, omega)
print(h10)
```

<!-- grade_id: cell-0a0af383ed825f38, points: 1, grade: true, solution: false -->
```python
assert np.allclose(h8, h9), "Output is not correct"
```

```python
print(omega_mat2)
```